Male rats emit aversive 44-kHz ultrasonic vocalizations during prolonged Pavlovian fear conditioning

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public Review):

The exclusive use of males is a major concern lacking adequate justification and should be disclosed in the title and abstract to ensure readers are aware of this limitation. With several reported sex differences in rat vocal behaviors this means caution should be exercised when generalizing from these findings. The occurrence of an estrus cycle in typical female rats is not justification for their exclusion. Note also that male rodents experience great variability in hormonal states as well, distinguishing between individuals and within individuals across time. The study of endocrinological influences on behavior can be separated from the study of said behavior itself, across all sexes. Similarly, concerns about needing to increase the number of animals when including all sexes are usually unwarranted (see Shansky [2019] and Phillips et al. [2023]).

As suggested by the Reviewer, we have disclosed the use of males in the title and the abstract. Also, we have added the statement that research on female rat subjects is required: “Here we are showing introductory evidence that 44-kHz vocalizations are a separate and behaviorally-relevant group of rat ultrasonic calls. These results require further confirmations and additional experiments, also in form of repetition, including research on female rat subjects.”

Regarding the analysis where calls were sorted using DBSCAN based on peak frequency and duration, my comment on the originally reviewed version stands. It seems that the calls are sorted by an (unbiased) algorithm into categories based on their frequency and duration, and because 44kHz calls differ by definition on frequency and duration the fact that the algorithm sorts them as a distinct category is not evidence that they are "new calls [that] form a separate, distinct group". I appreciate that the authors have softened their language regarding the novelty and distinctness of these calls, but the manuscript contains several instances where claims of novelty and specificity (e.g. the subtitle on line 193) is emphasized beyond what the data justifies.

We further softened our language regarding novelty and distinctness of 44-kHz vocalizations – including the aforementioned subtitle. However, in response, we would like to bring to the readers’ attention that all major groups of calls, i.e., long 22-kHz calls, short 22-kHz calls, and 50-kHz vocalization, are also defined in our manuscript and in the literature by their frequency and duration. However not one of these groups was identified separately by DBSCAN clustering excepting the 44kHz vocalizations. If they were not a distinct group, we would expect the 44-kHz and 50-kHz vocalizations to blend first (because of the similar frequencies) or 44-kHz and 22-kHz calls to merge first (because of the similar durations), but they do not in this unbiased examination.

The behavioral response to call playback is intriguing, although again more in line with the hypothesis that these are not a distinct type of call but merely represent expected variation in vocalization parameters. Across the board animals respond rather similarly to hearing 22 kHz calls as they do to hearing 44 kHz calls, with occasional shifts of 44 kHz call responses to an intermediate between appetitive and aversive calls. This does raise interesting questions about how, ethologically, animals may interpret such variation and integrate this interpretation in their responses. However, the categorical approach employed here does not address these questions fully.

This paragraph is exactly the same as in the previous review. There was no comment regarding our previous answer. Here is the previous answer:

“We are unsure of the Reviewer’s critique in this paragraph and will attempt to address it to the best of our understanding. Our finding of up to >19% of long seemingly aversive, 44-kHz calls, at a frequency in the define appetitive ultrasonic range (usually >32 kHz) is unexpected rather than “expected”. We would agree that aversive call variation is expected, but not in the appetitive frequency range.

Kindly note the findings by Saito et al. (2019), which claim that frequency band plays the main role in rat ultrasonic perception. It is possible that the higher peak frequency of 44-kHz calls may be a strong factor in their perception by rats, which is, however, modified by the longer duration and the lack of modulation.

Also, from our experience, it is quite challenging to demonstrate different behavioral responses of naïve rats to pre-recorded 22-kHz (aversive) vs. 50-kHz (appetitive) vocalizations. Therefore, to demonstrate a difference in response to two distinct, potentially aversive, calls, i.e., 22kHz vs. 44-kHz calls, to be even more difficult (as to our knowledge, a comparable experiment between short vs. long 22-kHz ultrasonic vocalizations, has not been done before).

Therefore, we do not take lightly the surprising and interesting finding that “animals respond rather similarly to hearing 22 kHz calls as they do to hearing 44 kHz calls, with occasional shifts of 44 kHz call responses to an intermediate between appetitive and aversive calls”. We would rather put this description in analogous words: “the rats responded similarly to hearing 44-kHz calls as they did to hearing aversive 22-kHz calls, especially regarding heart-rate change, despite the 44-kHz calls occupying the frequency band of appetitive 50-kHz vocalizations” and “other responses to 44-kHz calls were intermediate, they fell between response levels to appetitive vs. aversive playback” – which we added to the Discussion.

Finally, we acknowledge that our findings do not present a finite and complete picture of the discussed aspects of behavioral responses to the presented ultrasonic stimuli (44-kHz vocalizations). Therefore, we have incorporated the Reviewer’s suggestion in the discussion. The added sentence reads: “Overall, these initial results raise further questions about how, ethologically, animals may interpret the variation in hearing 22-kHz vs. 44-kHz calls and integrate this interpretation in their responses.”

I appreciate the amendment in discussing the idea of arousal being the key determinant for the increased emission of 44kHz, and the addition of other factors. Some of the items in this list, such as annoyance/anger and disgust/boredom, don't really seem to fit the data. I'm not sure I find the idea that rats become annoyed or disgusted during fear conditioning to be a particularly compelling argument. As such the list appears to be a collection of emotion-related words, with unclear potential associations with the 44kHz calls.

We agree that most of the factors listed are not supported by the data. These are hypotheses and speculations only – hence, an assumption / tentative statement, i.e., “It could also be argued that…”. We have changed it into “It could also be speculated that…”.

Later in the Discussion the authors argue that the 44kHz aversive calls signal an increased intensity of a negative valence emotional state. It is not clear how the presented arguments actually support this. For example, what does the elongation of fear conditioning to 10 trials have to do with increased negative emotionality? Is there data supporting this relationship between duration and emotion, outside anthropomorphism? Each of the 6 arguments presented seems quite distant from being able to support this conclusion.

We have added a description summarizing the literature that expounds the differences in employing one-two vs. five-ten foot-shocks during fear-conditioning training. It says:

“Importantly, it has been demonstrated multiple times that training rats with several electric foot-shocks (i.e., 5-10 shocks) produces a qualitatively different kind of fear-memory compared to training with only 1-2 shocks. Training with more numerous shocks has been shown to result in augmented freezing (e.g., Fanselow and Bolles, 1979, Haubrich et al., 2020, Haubrich and Nader, 2023, Poulos et al., 2016, Wang et al., 2009) which reflects a more intense fear-memory that is resistant to extinction (Haubrich et al., 2020, Haubrich and Nader, 2023), resistant to reconsolidation blockade (Haubrich et al., 2020, Wang et al., 2009, Finnie and Nader, 2020), associated with downregulation of NR2B NMDA-receptor subunits as well as elevated amyloid-beta concentrations in the lateral and basal amygdala (Finnie and Nader, 2020, Wang et al., 2009). Additionally, it involves activation of the noradrenaline-locus coeruleus system (Haubrich et al., 2020) and collective changes in connectivity across multiple brain regions within the neural network (Haubrich and Nader, 2023).

Notably, it has also been shown that higher freezing as a result of fear-conditioning training correlates with increased concentrations of stress hormone, corticosterone, in the blood (Dos Santos Correa et al., 2019). The rats subjected to 6- and 10-trial fear conditioning, whose results are reported herein (Tab. 1/Exp. 2/#7,8,11,12; n = 73), also demonstrated higher freezing than rats subjected to 1trial conditioning (Tab. 1/Exp. 2/#6,10; n = 33), which is reported elsewhere (Olszynski et al., 2021, Fig. S1C-E; Olszynski et al., 2022, Fig. S1D-G). Therefore, we postulate that emission of 44-kHz calls is associated with increased stress and the training regime forming robust memories.”

In sum, rather than describing the 44kHz long calls as a new call type, it may be more accurate to say that sometimes aversive calls can occur at frequencies above 22 kHz. Individual and situational variability in vocalization parameters seems to be expected, much more so than all members of a species strictly adhering to extremely non-variable behavioral outputs.

This paragraph is exactly the same as in the previous review. There was no comment regarding our previous answer. Here is the previous answer:

“The surprising fact that there are presumably aversive calls that are beyond the commonly applied thresholds, i.e., >32 kHz, while sharing some characteristics with 22-kHz calls, is the main finding of the current publication. Whether they be finally assigned as a new type, subtype, i.e. a separate category or become a supergroup of aversive calls with 22-kHz vocalizations is of secondary importance to be discussed with other researchers of the field of study.

However, we would argue – by showing a comparison – that 22-kHz calls occur at durations of <300 ms and also >300 ms, and are, usually, referred to in literature as short and long 22-kHz vocalizations, respectively (not introduced with a description that “sometimes 22-kHz calls can occur at durations below 300 ms”). These are then regarded and investigated as separate groups or classes usually referred to as two different “types” (e.g., Barker et al., 2010) or “subtypes” (e.g., Brudzynski, 2015). Analogously, 44-kHz vocalizations can also be regarded as a separate type or a subtype of 22kHz calls. The problem with the latter is that 22-kHz vocalizations are traditionally and predominantly defined by 18–32 kHz frequency bandwidth (Araya et al., 2020; Barroso et al., 2019; Browning et al., 2011; Brudzynski et al., 1993; Hinchcliffe et al., 2022; Willey & Spear, 2013).”

Reviewer #1 (Recommendations For The Authors):

Additional considerations:

Abstract: The 19.4% seems to be the percentage of 44 kHz calls observed during the 9th trial of the 10trial experiment, not the percentage of calls that were 44kHz during bouts of freezing.

We clarified the sentence. It now says:

“We observed 44-kHz calls to be associated with freezing behavior during fear conditioning training, during which they constituted up to 19.4% of all calls and most of them appeared next to each other forming groups of vocalizations (bouts).”

Abstract: "We hope that future investigations of 44-kHz calls in rat models of human diseases will contribute to expanding our understanding and therapeutic strategies related to human psychiatric conditions." This sounds like a far too strong of an implication provided the link between these calls and models of human psychiatric conditions is not clear.

We agree, the link is not clear. Therefore we only express our hope. We hope “the link” is there. While other ultrasonic calls are already being investigated in such animal models, training regimes employing numerous electric shocks are used as models of PTSD, helplessness etc.

Line 101: Seems a strong assumption to state the authors of the other publication were inspired by this paper, unless there is personal communication corroborating this.

The wording of the sentence has been changed.

It is still not clear why both Friedman and Wilcoxon tests were used, especially in situations where only one result seems to be referenced (for example on line 108-109).

We added the explanation within Methods: “In particular, the Friedman test was used to assess the presence of change within the sequence of several ITI, while the Wilcoxon test was used for the difference between the first and the last ITI analyzed.”

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Rats are believed to communicate their emotional state by emitting two distinct types of ultrasonic vocalizations. The first is long “22-kHz” vocalizations (>300 ms, <32 kHz) with constant frequency, signaling aversive states and the second, is short “50-kHz” calls (<150 ms, >32 kHz), often frequency-modulated, in appetitive situations. Here we are first to describe a new-type of aversive vocalizations emitted at a higher pitch by Wistar and spontaneously hypertensive rats (SHR) in an intensified aversive state – prolonged fear conditioning. These calls, which we named “44-kHz” vocalizations, are long (>150 ms), generally at a constant frequency (usually within 35-50 kHz range) and have an overall spectrographic image similar to 22-kHz calls. Some 44-kHz vocalizations are comprised of both 22-kHz-like and 44-kHz-like elements. Furthermore, two separate clustering methods confirmed that these 44-kHz calls are distinct from other vocalizations. We observed 44-kHz calls to be associated with freezing behavior during fear conditioning, during which they constituted up to 19.4% of all calls. We also show that some of rats’ responses to the playback of 44-kHz calls were more akin to that of aversive calls, e.g., heart rate changes, whereas other responses were at an intermediate level between aversive and appetitive calls. Our results suggest that rats have a wider vocal repertoire than previously believed, and current definitions of major call types may require reevaluation. We hope that future investigations of 44-kHz calls in rat models of human diseases will contribute to expanding our understanding and therapeutic strategies related to human psychiatric conditions.

Revised version: Rats are believed to communicate their emotional state by emitting two distinct types of ultrasonic vocalizations. The first is long “22-kHz” vocalizations (>300 ms, <32 kHz) with constant frequency, signaling aversive states and the second, is short “50-kHz” calls (<150 ms, >32 kHz), often frequency-modulated, in appetitive situations. Here we describe aversive vocalizations emitted at a higher pitch by male Wistar and spontaneously hypertensive rats (SHR) in an intensified aversive state – prolonged fear conditioning. These calls, which we named “44-kHz” vocalizations, are long (>150 ms), generally at a constant frequency (usually within 35-50 kHz range) and have an overall spectrographic image similar to 22-kHz calls. Some 44-kHz vocalizations are comprised of both 22-kHz-like and 44-kHz-like elements. Furthermore, two separate clustering methods confirmed that these 44-kHz calls can be separated from other vocalizations. We observed 44-kHz calls to be associated with freezing behavior during fear conditioning training, during which they constituted up to 19.4% of all calls and most of them appeared next to each other forming uniform groups of vocalizations (bouts). We also show that some of rats’ responses to the playback of 44-kHz calls were more akin to that of aversive calls, e.g., heart rate changes, whereas other responses were at an intermediate level between aversive and appetitive calls. Our results suggest that rats have a wider vocal repertoire than previously believed, and current definitions of major call types may require reevaluation. We hope that future investigations of 44-kHz calls in rat models of human diseases will contribute to expanding our understanding and therapeutic strategies related to human psychiatric conditions.

Introduction

Charles Darwin wrote: “That the pitch of the voice bears some relation to certain states of feeling is tolerably clear” (Darwin, 1872). This has also been tolerably clearly observed and widely described for ultrasonic vocalizations of rats (Brudzynski, 2019, Brudzynski, 2021, Simola and Granon, 2019) which emit low-pitched aversive calls and high-pitched appetitive calls. The former are “22-kHz” vocalizations (Figs 1A, 2A), with 18 to 32 kHz frequency range, monotonous and long, usually >300 ms, and are uttered in distress (Brudzynski, 2013, Brudzynski, 2019, Brudzynski, 2021, Simola and Granon, 2019). The latter are “50-kHz” vocalizations (Fig. 1C), are relatively short (10-150 ms), frequency-modulated, usually within 35-80 kHz, and they signal appetitive and rewarding states (Simola and Granon, 2019, Brudzynski, 2013, Brudzynski, 2019, Brudzynski, 2021). Therefore, these two types of calls communicate the animal’s emotional state to their social group (Brudzynski, 2013). Low-pitch (<32 kHz), short (<300 ms; Fig. 1B) calls, assumed to also express a negative aversive state, have been described but their role is not clearly established (Brudzynski, 2013). Notably, high-pitch (>32 kHz), long and monotonous ultrasonic vocalizations have not yet been described. Here we show these unmodulated rat vocalizations with peak frequency1 at about 44 kHz (Figs 1B, 1E, 2B), emitted in aversive experimental situations, especially in prolonged fear conditioning.

Characteristics of vocalizations emitted by Wistar rats during fear conditioning with ten aversive foot-shocks

(Tab. 1/ Exp. 1-3/#2,4,8,13; n = 46). A – some rats produced aversive 22-kHz vocalizations with typical features, i.e., constant-frequency of <32 kHz, >300 ms duration – both values marked as dotted lines); example emission from one rat. B – some rats produced 44-kHz vocalizations with constant frequency of >32 kHz and long duration (>150 ms); example emission from one rat. C – rats which emitted aversive vocalizations during fear session, produced 50-kHz vocalizations during appetitive playback session the following day (full data published in Olszyński et al., 2021); representative data from same rat in A. D – the onset of long 22-kHz alarm calls typically occurred after first shock stimulus (vertical dotted lines mark time of shock deliveries in DE); note the gradual rise in peak frequencya, not exceeding 32 kHz (horizontal dotted line in DE); data from the same rat as AC. E – in rats that emitted 44-kHz calls, the onset was usually delayed to after several foot-shocks; note the gradual rise in peak frequency of both long 22-kHz and 44-kHz vocalizations throughout training (comp. Fig. 1S2CD); data from same rat in B). F – call rate of long 22-kHz calls was higher than 44-kHz calls (*p < 0.05, **p < 0.01, ***p < 0.001) and with different time-course – maximum number of 22-kHz calls at ITI-3 (higher than ITI-1, 2, 5-10; <0.0001–0.0005 p levels); and higher number of 44-kHz calls at ITI-5-10, i.e., 6.6 ± 2.3 vs. ITI-1-4, i.e., 0.4 ± 0.2; p < 0.0001; all Wilcoxon); numbers of ITI (inter-trial-intervals) correspond to the numbers of previous foot-shocks, values are means ± SEM. G – long 44-kHz vocalizations had a higher incidence rate (15.5%) than short 22-kHz (8.8%) and 50-kHz calls (5.6%); values are calculated for sum of all vocalizations obtained during entire training sessions (there were fewer 50-kHz calls, i.e., 3.7%, when vocalizations prior to the first shock were not included). A-E: dots reflect specified single rat values. FG: n = 46, other results from these rats are previously published (Olszyński et al., 2021, Olszyński et al., 2022).

Five subtypes (B-F) of high frequency 44-kHz aversive vocalizations.

A – standard aversive 22-kHz vocalization with peak frequency <32 kHz (peak frequency = 24.4 kHz). 44-kHz aversive vocalization subtypes: B – flat (constant frequency call; peak frequency = 42.4 kHz), C – step up (peak frequency = 39.5 kHz), D – step down (peak frequency = 52.2 kHz), E – insert (peak frequency = 38.5 kHz), F – complex (peak frequency = 46.3 kHz). G – percentage share of 44-kHz call-subtypes in all cases of detected 44-kHz vocalizations.

Results

Earlier version: New calls are high, long, unmodulated

Revised version: High, long, and unmodulated calls

In three separate experiments (all summarized in Tab. 1/Exp.1-3, see Methods), i.e., one with trace-fear-conditioning (Tab. 1/Exp. 1) and two with delay-fear-conditioning (Tab. 1/Exp. 2-3), one of which has already been described (Tab. 1/Exp. 2, Olszyński et al., 2021, Olszyński et al., 2022), 53 of all 84 conditioned Wistar rats (Tab. 1/Exp. 1-3/#2,4,6-8,13, Figs 1B, 1E, 1S1BC) displayed vocalizations that were high-pitched, i.e., in the range of 50-kHz calls, but long and monotonous (Fig. 2B). These vocalizations, e.g., top-right group in Figs 1B and 1S1C, were outside the defined range (Brudzynski, 2019, Brudzynski, 2021, Simola and Granon, 2019) for both 50-kHz (bottom-right group in Figs 1C, 1S1A-C) and 22-kHz calls (top-left group in Figs 1A, 1B, 1S1A-C). These vocalizations were also observed in a different rat strain acquired from a different breeding colony, i.e., spontaneously hypertensive rats (SHR) (Okamoto and Aoki, 1963), also trained in delay fear conditioning (Tab. 1/Exp. 2/#10-12; Olszyński et al., 2022). Six of the 49 conditioned SHR displayed high-pitch, long, monotonous vocalizations (e.g., Fig. 2S1G); moreover, we observed more of these vocalizations in Wistar rats compared to SHR (Tab. 1/Exp. 2/#6-8,10-12) in both training, p < 0.0001, and test sessions, p = 0.0030, Mann-Whitney.

Earlier version: Overall, we analyzed 140,149 vocalizations from all fear conditioning experiments (Tab. 1/Exp. 1-3/#1-13, n = 218) and through trial-and-error, we set new criteria, namely peak frequency of >32 kHz and >150 ms duration to define the new-type calls. We manually verified the results on the spectrogram using these parameters and only 308 calls (0.2%) were incorrectly assigned (i.e., exceptionally long 50-kHz vocalizations misplaced in the new-type group or borderline-short vocalizations of the new-type misplaced to 50-kHz calls). Hence the new parameters correctly assigned 99.8% of cases and are thus effective to distinguish the new-type calls in an automated fashion. Finally, 10,445 new-type calls were identified, which constituted 7.5% of the total calls during fear conditioning experiments (Tab. 1/Exp. 1-3; comp. Fig. 1G). These vocalizations have a peak frequency range from 32.2 to 51.5 kHz (95% of cases) with an average peak frequency of 42.1 kHz, and they exhibited 43.8 kHz peak frequency at the cluster center in a DBSCAN analysis (Fig. 3A). In line with the accepted nomenclature convention, underlining the relationship with 22-kHz vocalizations, we christened this new-type of ultrasonic calls as “44-kHz vocalizations”.

Revised version: Overall, we analyzed 140,149 vocalizations from all fear conditioning experiments (Tab. 1/Exp. 1-3/#1-13, n = 218) and through trial-and-error, we set new criteria, namely peak frequency of >32 kHz and >150 ms duration to define the calls described above. We manually verified the results on the spectrogram using these parameters and only 308 calls (0.2%) were incorrectly assigned (i.e., exceptionally long 50-kHz vocalizations misplaced in the newly-defined group or borderline-short vocalizations of the newly-defined group misplaced to 50-kHz calls). Hence the new parameters correctly assigned 99.8% of cases and are thus effective to distinguish the newly-defined calls in an automated fashion. Finally, 10,445 newly-defined calls were identified, which constituted 7.5% of the total calls during fear conditioning experiments (Tab. 1/Exp. 1-3; comp. Fig. 1G). These vocalizations have a peak frequency range from 32.2 to 51.5 kHz (95% of cases) with an average peak frequency of 42.1 kHz, and they exhibited 43.8 kHz peak frequency at the cluster center in a DBSCAN analysis (Fig. 3A). In line with the accepted nomenclature convention, underlining the relationship with 22-kHz vocalizations, we christened this newly-defined calls as “44-kHz vocalizations”.

Clustering of ultrasonic vocalizations from fear conditioning sessions using two independent methods.

A – DBSCAN algorithm (ε = 0.14) clustering of vocalizations from all fear conditioning experiments (Tab. 1/Exp. 1-3/#1-13, n = 218), silhouette coefficient = 0.198, two clusters emerge, cluster of green dots n = 77,243 (due to high generality of cluster average peak frequency and duration deemed redundant), cluster of red dots n = 5,646 (average peak frequency = 43,826.6 Hz, average duration = 0.524 s), some calls were not assigned to any cluster, i.e., outlier vocalizations, black dots, n = 4,139. BC – clustering by k-means algorithm and visualization of calls emitted by selected rats, i.e., with >30 of 44-kHz vocalizations, during trace and delay fear conditioning training (n = 26, selected from Tab. 1/Exp. 1-3/#2,4,7,8,11-13), total number of calls n = 40,084. B – topological plot of ultrasonic calls using UMAP embedding, particular agglomerations of calls labeled with their type or subtype. C – spectrogram images from DeepSqueak software superposed over plot B, colors denote clusters from unsupervised clustering, number of clusters set using elbow optimization (max number = 4), two clusters emerge; see also Fig. 3S1.

44-kHz calls in long aversive stimulation

We found 44-kHz vocalizations especially in rats which received multiple electric shocks. When we analyzed all Wistar rats that had undergone 10 trials of fear conditioning (Tab. 1/Exp. 1-3/#2,4,8,13; n = 46), these vocalizations were less frequent following the first trial (1.2 ± 0.4% of all calls), and increased in subsequent trials, particularly after the 5th (8.8 ± 2.8%), through the 9th (19.4 ± 5.5%, the highest value), to the 10th (15.5 ± 4.9%) trials, where 44-kHz calls gradually replaced 22-kHz vocalizations in some rats (Fig. 1F, 1S2AB, Video 1; comp Fig. 1D vs. 1E). Please note, majority of the 22-kHz calls were emitted after the 3rd shock, i.e., during the 3rd ITI (inter-trial-interval), while 44-kHz vocalizations were emitted in the second part of the training, i.e., 5th to 10th ITI (Fig. 1F, comp. Fig. 1S2AB). From this group of rats (n = 46), n = 41 (89.1%) emitted long 22-kHz calls, and 32 of them (69.6%) emitted 44-kHz calls, i.e., every animal that produced 44-kHz calls also emitted long 22-kHz calls (Fig. 1S2AB). The prevalence of 44-kHz calls varied greatly among individual rats, such that for n = 3 rats, 44-kHz vocalizations accounted for >95% of all calls during at least one ITI (e.g., 140 of total 142, 222 of 231, and 263 of 265 tallied 44-kHz calls), and in n = 9 rats, 44-kHz vocalizations constituted >50% of calls in more than one ITI. The prevalence of 44-kHz calls in all experimental conditions analyzed in all animal groups is shown in Fig. 1S3.

Notably, there were more 44-kHz vocalizations during fear conditioning training than testing in all fear-conditioned Wistar rats (Tab. 1/Exp. 1-3/#2,4,6-8,13; n = 84; 3.63 ± 0.99 vs. 0.23 ± 0.13 calls/min; p < 0.0001; Wilcoxon).

Earlier version: In a recent publication during this paper’s review process, Gonzalez-Palomares et al. (2023), inspired by our current findings, investigated and reported 44-kHz vocalizations following prolonged (10-trial procedure) odor fear conditioning. These calls were observed predominantly during the late ITI, i.e., 8th-10th ITI (Gonzalez-Palomares et al., 2023; Fig. S4C; please note 4th-7th ITI were not investigated) after the shock presentations (Fig. S4B therein), which complement our results.

Revised version: In a recent publication during this paper’s review process, Gonzalez-Palomares et al. (2023), in line with the findings reported here, investigated and described 44-kHz vocalizations following prolonged (10-trial protocol) odor fear conditioning. These calls were observed predominantly during the late ITI, i.e., 8th-10th ITI (Gonzalez-Palomares et al., 2023; Fig. S4C; please note 4th-7th ITI were not investigated) after the shock presentations (Fig. S4B therein), which complement our results.

Changes in frequency, duration, and mean power of long aversive calls during conditioning

Analyzing Wistar rats that undergone 10 trials of fear conditioning (Tab. 1/Exp. 1-3/#2,4,8,13; n = 46), we also observed the frequencies of 22-kHz calls to gradually rise throughout fear conditioning training, i.e., during subsequent ITI – from 24.5 ± 0.1 to 27.9 ± 0.4 kHz (Figs 1DE, 1S2C; p < 0.0001, Friedman, p = 0.0039, Wilcoxon). The frequency levels of 44-kHz vocalizations also appeared to rise – from 37.8 ± 2.1 to 39.6 ± 1.3 kHz (Figs 1E, 1S2C) but we were unable to statistically demonstrate it (p = 0.0155, Friedman, p = 0.0977, Wilcoxon).

There was a shortening of long 22-kHz calls during the first four ITI from 969.6 ± 43.1 ms to 794.6 ± 39.8 ms (p < 0.0001, Friedman; p < 0.0001, Wilcoxon, Fig. 1S2D), while 44-kHz vocalizations were longest during the 4th ITI (the time of their substantial appearance, comp. Fig. 1F), i.e., 775.0 ± 135.7 ms, and shortened over subsequent ITI (619.6 ± 58.1 ms for the 10th ITI, Fig. 1S2D, p = 0.0227, Friedman; p = 0.0234, Wilcoxon).

Finally, the sound mean power of 44-kHz vocalizations appeared to remain stable throughout the 10-trial sessions, while during the first half of the training, i.e., 1st-5th ITI, 22-kHz calls were not only significantly more frequent but also louder than during the second half, i.e., 6th-10th ITI (p < 0.0001, Wilcoxon). Consequently, long 22-kHz calls appeared louder than 44-kHz calls (p = 0.0397-0.0038, Mann-Whitney). However, in the second half of the session, this difference dissipated due to the diminishing amplitude of 22-kHz vocalizations (p = 0.0083, Friedman; p = 0.0046, Wilcoxon), while the amplitude of 44-kHz calls remained stable (p = 0.0663, Friedman; p = 0.2661, Wilcoxon; 6th ITI through 10th ITI for both; Fig. 1S2E). After adjusting for angle-dependent hardware attenuation (see Methods, Sound mean power), the situation reversed (Fig. 1S2F). Both long 22-kHz and 44-kHz vocalizations showed similar amplitude levels during the first half of the fear conditioning session, while during the 6th-10th ITI, 44-kHz calls were significantly louder than long 22-kHz calls (p = 0.0007-0.0097, Mann-Whitney).

44-kHz calls linked to freezing

We investigated the freezing behavior of all Wistar rats emitting 44-kHz vocalizations during 10 trials of fear conditioning (Tab. 1/Exp. 1-3/#2,4,8,13; n = 46). The training sessions were divided into 10-s-long time bins, from which we analyzed only the bins that had exclusively long 22-kHz or 44-kHz calls. For comparison, we also measured the freezing levels during the first 5 min of the trial (baseline freezing levels before any foot-shocks) as well as the bins in which animals did not vocalize (from the period after the 1st shock to the end of the session). Of the n = 46 rats analyzed, n = 41 emitted 22-kHz vocalizations, from which n = 32 also emitted 44-kHz vocalizations, from which only n = 21 were determined to have both – 10-s-long bins of 22-kHz calls only and 44-kHz calls only (Tab. 2A). Freezing during the bins of 22-kHz calls only (p < 0.0001, for both groups) and during 44-kHz calls only bins (p = 0.0003) was higher than during the first 5 min baseline freezing levels of the session. Also, the freezing associated with emissions of 44-kHz calls only was higher than during bins with no ultrasonic vocalizations (p = 0.0353), and it was also 9.9 percentage points higher than during time bins with only long 22-kHz vocalizations, but the difference was not significant (p = 0.1907; all Wilcoxon).

To further investigate this potential difference, we measured freezing during the emission of randomly selected single 44-kHz and 22-kHz vocalizations. The minimal freezing behavior detection window was reduced to compensate for the higher resolution of the measurements (3, 5, 10, or 15 video frames were used). There was no difference in freezing during the emission of 44-kHz vs. 22-kHz vocalizations for ≥150-ms-long calls (3 frames, p = 0.2054) and for ≥500-ms-long calls (5 frames, p = 0.2404; 10 frames, p = 0.4498; 15 frames, p = 0.7776; all Wilcoxon, Tab. 2B).

44-kHz calls sorted into five subtypes

While the majority of 44-kHz vocalizations were of continuous unmodulated frequency (Fig. 2B), some comprised additional elements. Based on the composition of individual call elements and their relation to each other, we manually sorted the calls into five categories (Fig. 2B-F). If the start (prefix) or end (suffix) portion of a call was less than 1/5th the length of the following or previous element, this portion of the call was not considered in its categorization into the five subtypes. The names and descriptions of the five subtypes are: flat – single element with near constant frequency and little to no interruptions to the sound continuity on the spectrogram; step up – two elements with an instantaneous frequency jump, where the first element is of lower frequency; step down – two elements with an instantaneous frequency jump, where the first element is of higher frequency; insert – three elements with an instantaneous frequency change, where the middle element is of different frequency; complex – more than three elements with instantaneous frequency changes.

44-kHz and 22-kHz calls closely related

44-kHz were emitted in aversive behavioral situations – as 22-kHz calls are observed (Antoniadis and McDonald, 1999, Dupin et al., 2019, Taylor et al., 2017). Both types of calls are long (usually >300 ms) and frequency-unmodulated. Some of the elements constituting such as step up; step down; insert and complex 44-kHz vocalizations (Fig. 2C-F) were at a lower frequency – typical for 22-kHz vocalizations. Vice versa we also observed 22-kHz calls with 44-kHz-like elements. Therefore, we propose that these long 22-kHz and 44-kHz vocalizations constitute a supertype group of long unmodulated aversive calls (“long 22/44-kHz vocalizations”).

We observed a stable, approximately 1.5 ratio in peak frequency levels between 22-kHz and 44-kHz vocalizations within individual rats. Specifically, in fourteen rats (13 Wistar and 1 SHR) with a clear transition from 22-kHz to 44-kHz calls during the fear conditioning session (n = 14, selected from Tab. 1/Exp. 1-3/#2,4,6-8,10-13), the proportion between the frequencies of the long 22-kHz vocalizations and the long 44-kHz calls was 1.48 ± 0.02. Similar results were obtained for 70 step up (1.53 ± 0.03) and 65 step down (1.59 ± 0.02) 44-kHz calls – altogether suggesting a 1.5-times or 3:2 frequency ratio. This ratio and its relevance has been observed in invertebrates and vertebrates including human speech and music (Hoeschele, 2017). In music theory, 3:2 frequency ratio is referred to as a perfect fifth and is often featured, e.g., the first two notes of the Star Wars 1977 movie (ascending, i.e., step up; comp. Fig. 2C, Track 1) and Game of Thrones 2011 television series (descending, i.e., step down; comp. Fig. 2D, Track 2) theme songs. All of which may point to a common basis for this sound interval and its prevalence which could be explained by the observation that all physical objects capable of producing tonal sounds generate harmonic vibrations, the most prominent being the octave, perfect fifth, and major third (Christensen, 1993, discussed in Bowling and Purves, 2015).

Added in the revised version: 44-kHz calls separated in cluster analyses

Removed in the revised version: New calls form separate, distinct group

Earlier version: Next, we showed that 44-kHz calls indeed constitute a distinct, separate type of ultrasonic vocalizations as it was sorted into isolated clusters by two different methods. First, using the DBSCAN algorithm method based on calls’ peak frequency and duration, we were able to divide all vocalizations recorded during all training sessions into 44-kHz vocalizations vs. all other vocalizations as two separate clusters (Fig. 3A). Secondly, a clustering algorithm that includes call contours, i.e., k-means with UMAP projection done via DeepSqueak (Figs 3BC, 3S1), sorted 44-kHz vocalizations of different subtypes including unusual ones (Fig. 2S1A-F), into topologically-separate groups. Notably, flat 44-kHz calls were consistently in a separate cluster from 22-kHz calls Figs 3C, 3S1B).

Revised version: Next, we showed that 44-kHz calls indeed constitute a separable type of ultrasonic vocalizations as it was sorted into isolated clusters by two different methods. First, using the DBSCAN algorithm method based on calls’ peak frequency and duration, we were able to divide all vocalizations recorded during all training sessions into 44-kHz vocalizations vs. all other vocalizations as two separate clusters (Fig. 3A). Secondly, a clustering algorithm that includes call contours, i.e., k-means with UMAP projection done via DeepSqueak (Figs 3BC, 3S1), sorted 44-kHz vocalizations of different subtypes including unusual ones (Fig. 2S1A-F), into topologically-separate groups. Notably, flat 44-kHz calls were consistently in a separate cluster from 22-kHz calls Figs 3C, 3S1B).

Added in the revised version: 44-kHz calls formed separate groups

Added in the revised version: We have examined transition probabilities for call-type transitions within all vocalizations and within vocalization bouts in all Wistar rats during 10 trials of fear conditioning training (Tab. 1/Exp. 1-3/#2,4,8,13; n = 46). Transition probabilities between vocalization types were defined by counting the number of specific pair sequences and dividing by all observed sequence pairs where the first type of signal is followed by any signal. The most probable call following long 22-kHz vocalization was another long 22-kHz vocalization (94.4% probability). Similarly, the most probable call following 44-kHz vocalization was another 44-kHz call (83.1%). These values augmented when the analysis was limited to vocalizations emitted only in bouts, i.e., with <320 ms time gaps between calls (Wöhr et al., 2005), and reached 96.3% and 86.5%, respectively. Examples of parts of groups of both types of calls are demonstrated in Fig. 3S2.

Earlier version: Specific response to 44-kHz playback

Revised version: Response to 44-kHz playback

To describe the behavioral and physiological impact of 44-kHz vocalizations, we performed playback experiments in two separate groups of rats (Methods, Figs 4, 4S1). Overall, the responses to 44-kHz aversive calls presented from the speaker were either similar to 22-kHz vocalizations or in-between responses to 22-kHz and 50-kHz playbacks. For example, the heart rate of rats exposed to 22-kHz and 44-kHz vocalizations decreased, and increased to 50-kHz calls (Fig. 4A, comp. Olszyński et al., 2020). Whereas the number of vocalizations emitted by rats was highest during and after the playback of 50-kHz, intermediate to 44-kHz and lowest to 22-kHz playbacks (Figs 4BC, 4S1EF). Additionally, the duration of 50-kHz vocalizations emitted in response to 44-kHz playback was also intermediate, i.e., longer than following 22-kHz playback (Fig. 4D) and shorter than following 50-kHz playback (Figs 4D, 4S1G). Finally, similar tendencies were observed in the distance travelled and time spent in the half of the cage adjacent to the speaker (Fig. 4S1A-D).

Physiological and behavioral response to playback of 44-kHz calls (vs. 50-kHz and 22-kHz calls) presented from a speaker to naïve Wistar rats.

A – heart rate (HR); B – the number of emitted vocalizations. AB – gray sections correspond to the 10-s-long ultrasonic playback. Each point is a mean for a 10-s-long time-interval with SEM. CD – properties of 50-kHz vocalizations emitted in response to ultrasonic playback, i.e., number of calls (C) and duration (D) calculated from the 0-120 s range. A – 50-kHz playback resulted in HR increase (playback time-interval vs. 10-30 s time-interval, p = 0.0007), while the presentation of the aversive playbacks resulted in HR decrease, both in case of 22-kHz (p < 0.0001) and 44-kHz (p = 0.0014, average from -30 to -10 time-intervals (i.e., “before”) vs. playback interval, all Wilcoxon), which resulted in different HR values following different playbacks, especially at +10 s (p = 0.0097 for 50 kHz vs. 22-kHz playback; p = 0.0275 for 50 kHz vs. 44-kHz playback) and +20 s time-intervals (p = 0.0068, p = 0.0097, respectively, all Mann-Whitney). B – 50-kHz playback resulted also in a rise of evoked vocalizations (before vs. 10-30 s time-interval, p = 0.0002, Wilcoxon) as was the case with 44-kHz playback (p = 0.0176 in respective comparison), while no rise was observed following 22-kHz playback (p = 0.1777). However, since the increase in vocalization was robust in case of 50-kHz playback, the number of emitted vocalizations was higher than after 22-kHz playback (e.g., p < 0.0001 during 0-30 time-intervals) as well as after 44-kHz playback (e.g., p < 0.0001 during 0-10 time-intervals, both Mann-Whitney). Finally, when the increases in the number of emitted ultrasonic calls in comparison with before intervals were analyzed, there was a difference following 44-kHz vs. 22-kHz playbacks during 30 s and 40 s time intervals (p = 0.0420 and 0.0430, respectively, Wilcoxon). C – During the 2 min following the onset of the playbacks, rats emitted more ultrasonic calls during and after 50-kHz playback in comparison with 22-kHz (p < 0.0001) and 44-kHz (p = 0.0011) playbacks. The difference between the effects of 22-kHz and 44-kHz playbacks was not significant (p = 0.2725, comp. Fig. 4S1F; all Mann-Whitney). D – Ultrasonic 50-kHz calls emitted in response to playback differed in their duration, i.e., they were longer to 50-kHz (p = 0.0004) and 44-kHz (p = 0.0273, both Mann-Whitney) playbacks than to 22-kHz playback. * 50-kHz vs. 44-kHz, $ 50-kHz vs. 22-kHz, # 22-kHz vs. 44-kHz; one character (*, $ or #), p < 0.05; two, p < 0.01; three, p < 0.001; Mann-Whitney (AB) or Wilcoxon (CD). Values are means ± SEM, n = 13-16.

Discussion

As Charles Darwin noted above (Darwin, 1872) and other researchers have confirmed (Briefer et al., 2012), the frequency level of animal calls is a vocal parameter that changes in accordance with its arousal state (intensity) or emotional valence (positive/negative state). The frequency shifts towards both higher and lower levels, i.e., alterations were observed during both positive (appetitive) and negative (agonistic/aversive) situations, however, as a general rule, frequency usually increases with an increase in arousal (Briefer et al., 2012). We would like to propose a hypothesis that our prolonged fear conditioning increased the arousal of the rats with no change in the valence of the aversive stimuli.

Earlier version: It could also be argued that several factors, apart from increased arousal, contributed to the emergence of 44-kHz vocalizations in our fear-conditioned rats, e.g., heightened fear, stress/anxiety, annoyance/anger, disgust/boredom, grief/sadness, despair/helplessness, and weariness/fatigue. It is not possible, at this stage, to definitively determine which factors played a decisive role. Please note that the potential contribution of these factors is not mutually exclusive.

Revised version: It could also be speculated that several other factors, apart from increased arousal, contributed to the emergence of 44-kHz vocalizations in our fear-conditioned rats, e.g., heightened fear, stress/anxiety, annoyance/anger, disgust/boredom, grief/sadness, despair/helplessness, and weariness/fatigue. It is not possible, at this stage, to definitively determine which factors played a decisive role. Please note that the potential contribution of these factors is not mutually exclusive.

However, several arguments support the idea that 44-kHz vocalizations communicate an increased negative emotional state. First, in general, ultrasonic vocalizations serve as a means of communicating rats’ emotional state (Brudzynski, 2013). Second, the changing of the pitch of the voice bears some relation to certain states of feeling (Darwin, 1872). Third, 44-kHz calls were notably more frequent during prolonged aversive stimulation, i.e., the 5th-10th trials of fear conditioning. Fourth, they were linked to freezing. Fifth, they appeared as partial replacements of, established as aversive, 22-kHz calls – in the presence of the same painful stimulus. Sixth, numerous instances of vocalizations featured both 22-kHz-like and 44-kHz-like call-elements.

Also, several observations contradict the potential contribution of fatigue. The sound mean power of 44-kHz vocalizations was comparable to, or possibly even higher than, that of 22-kHz calls, despite the higher energy costs associated with producing higher-pitched calls (Sonninen and Hurme, 1998), i.e., the rats emitting 44-kHz calls invested additional energy to communicate their emotional state; both in vivo measurements (Riede, 2013) and computer modelling (Hakansson et al., 2022) demonstrated that producing calls of higher frequency, such as 50 kHz vs. 22 kHz, requires increased activity of various muscles. Additionally, the mean power of 44-kHz vocalizations remained strong and stable for several trials – in contrast to 22-kHz vocalizations. Finally, when 44-kHz calls started to appear in significant numbers, i.e., after the 4th-5th trials of fear conditioning, they were as long as 22-kHz vocalizations.

Concerning the latter, we observed a significant decrease in the mean power of 22-kHz vocalizations during the fear conditioning session. Such reduction could potentially be attributed to fatigue (as observed in humans, Kitch and Oates, 1994), despair (e.g., as a reaction to the lack of effects from repeated emissions of 22-kHz calls), or both. The reduction in the amplitude of 22-kHz calls during the 10-trial fear conditioning was also recently observed by others (Gonzalez-Palomares et al., 2023).

Added in the revised version: Importantly, it has been demonstrated multiple times that training rats with several electric foot-shocks (i.e., 5-10 shocks) produces a qualitatively different kind of fear-memory compared to training with only 1-2 shocks. Training with more numerous shocks has been shown to result in augmented freezing (e.g., Fanselow and Bolles, 1979, Haubrich et al., 2020, Haubrich and Nader, 2023, Poulos et al., 2016, Wang et al., 2009) which reflects a more intense fear-memory that is resistant to extinction (Haubrich et al., 2020, Haubrich and Nader, 2023), resistant to reconsolidation blockade (Haubrich et al., 2020, Wang et al., 2009, Finnie and Nader, 2020), associated with downregulation of NR2B NMDA-receptor subunits as well as elevated amyloid-beta concentrations in the lateral and basal amygdala (Finnie and Nader, 2020, Wang et al., 2009). Additionally, it involves activation of the noradrenaline-locus coeruleus system (Haubrich et al., 2020) and collective changes in connectivity across multiple brain regions within the neural network (Haubrich and Nader, 2023).

Added in the revised version: Notably, it has also been shown that higher freezing as a result of fear-conditioning training correlates with increased concentrations of stress hormone, corticosterone, in the blood (Dos Santos Correa et al., 2019). The rats subjected to 6- and 10-trial fear conditioning, whose results are reported herein (Tab. 1/Exp. 2/#7,8,11,12; n = 73), also demonstrated higher freezing than rats subjected to 1-trial conditioning (Tab. 1/Exp. 2/#6,10; n = 33), which is reported elsewhere (Olszynski et al., 2021, Fig. S1C-E; Olszynski et al., 2022, Fig. S1D-G). Therefore, we postulate that emission of 44-kHz calls is associated with increased stress and the training regime forming robust memories.

Amounting research points to the utility of rat ultrasonic vocalizations to alter emotional states, evidenced by behavioral changes, in tested rats via playback of affectively valenced calls (Bonauto et al., 2023). We have exposed rats to 44-kHz playback along with 22-kHz and 50-kHz playback. The experimental design (see methods for details) allowed us to compare rats’ responses to 22-kHz vs. 44-kHz playbacks especially – with 50-kHz playback used as a form of control or baseline. In general, the rats responded similarly to hearing 44-kHz calls as they did to hearing aversive 22-kHz calls, especially regarding heart-rate change, despite the 44-kHz calls occupying the frequency band of appetitive 50-kHz vocalizations. This is contrary to some observations (Saito et al., 2019) which suggested that frequency band plays the main role in rat ultrasound perception. Please factor in potential carry-over effects (resulting from hearing playbacks of the same valence in a row) in the differences between responses to 50-kHz vs. 22/44-kHz playbacks, especially, those observed before the signal (Fig. 4AB). Other responses to 44-kHz calls were intermediate, they fell between response levels to appetitive vs. aversive playback, which might signify some behavioral specificity and importance (or possibly confusion). These latter effects were similar in both playback experiments despite an array of methodological differences between them. Overall, these initial results raise further questions about how, ethologically, animals may interpret the variation in hearing 22-kHz vs. 44-kHz calls and integrate this interpretation in their responses.

The question also is, why have the 44-kHz vocalizations been overlooked until now? On one hand, long (or not that long as in Biały et al., 2019), frequency-stable high-pitch vocalizations have been reported before (e.g., Sales, 1979; Shimoju et al., 2020), notably as caused by intense cholinergic stimulation (Brudzynski and Bihari, 1990) or higher shock-dose fear conditioning (Wöhr et al., 2005). However, they have not been systematically defined, described, fully shown or demonstrated to be a separate type of vocalization. On the other hand, 44-kHz calls were likely omitted as the analyses were restricted to canonical groups, i.e. flat 22-kHz and short 50-kHz calls, with a sharp dividing frequency border between the two (e.g., Kalamari et al., 2021, Potasiewicz et al., 2020, Turner et al., 2019) or even a frequency ‘safety gap’ between 22-kHz and 50-kHz vocalizations (e.g., Silkstone and Brudzynski, 2019, Garcia et al., 2015). Moreover – many older bat-detectors had limited frequency-range detection (e.g., up to 40 kHz in Sales, 1991), when stress-evoked types of ultrasonic calls were being established. Finally, 44-kHz vocalizations are emitted much fewer than 22-kHz calls (Fig. 1FG).

Here we present introductory evidence that 44-kHz vocalizations are a separate and behaviorally-relevant group of rat ultrasonic calls. These results require further confirmations and additional experiments, also in form of replication, including research on female rat subjects. However, our results bring to awareness that rats employ these previously unrecognized, long, high-pitched and flat aversive calls in their vocal repertoire. Researchers investigating rat ultrasonic vocalizations should be aware of their potential presence and to not rely fully on automated detection of high vs. low-pitch calls.

Materials and methods

Animals

Wistar rats (n = 167) were obtained from The Center for Experimental Medicine of the Medical University of Bialystok, Poland; spontaneously hypertensive rats (SHR, n = 80) and Sprague-Dawley rats (n = 16) were from Mossakowski Medical Research Institute, Polish Academy of Sciences, Poland. All rats were males, 7 weeks of age on arrival, randomly assigned into groups and cage pairs where appropriate; housed with a 12 h light-dark cycle, ambient temperature (22–25 °C) with standard chow and water provided ad libitum. The animals were left undisturbed for at least one week before any procedures, then handled at least four times for 2 min by each experimenter directly involved for one to two weeks. All procedures were approved by Local Ethical Committees for Animal Experimentation in Warsaw.

Animal details: groups of animals used

Trace fear conditioning experiment

Wistar rats, both single-housed (n = 14) and pair-housed (n = 20), were implanted with radiotelemetric transmitters for measuring heart rate in an ultrasonic vocalization playback experiment previously described by us (Olszyński et al., 2020) after which, at 13 weeks of age, half of them (n = 17) were fear-conditioned (10 shocks), while the other half (n = 17) served as controls (Tab. 1/Exp. 1/#1-4, n = 34).

Delay fear conditioning experiment, rats with transmitters

Wistar rats (n = 94) and SHR (n = 80) were implanted with a radiotelemetric transmitters one week before fear conditioning during which they received 0, 1, 6 or 10 shocks at 12 weeks of age (Tab. 1/Exp. 2/#5-12, n = 174). All the details are described in Olszyński et al. (2021) and Olszyński et al. (2022).

Delay fear conditioning experiment, rats without transmitters

Wistar rats were housed in pairs; were not implanted with radiotelemetric transmitters to eliminate the potential effect of surgical intervention on vocalization; they received 10 conditioning stimuli at 12 weeks of age (Tab. 1/Exp. 3/#13, n = 10) – same as in Olszyński et al. (2021) and Olszyński et al. (2022).

Playback experiment, rats with transmitters

Wistar rats (n = 29) were housed in pairs; all were implanted with a radiotelemetric transmitter one week before the playback experiment. At 12 weeks of age, one group (n = 13) heard 50-kHz appetitive vocalization playback while the other (n = 16) 22-kHz and 44-kHz aversive calls (for details see below).

Playback experiment, rats without transmitters

Sprague Dawley rats (n = 16) were housed in pairs, were not implanted with the transmitters, and received 22-kHz, 44-kHz, and 50-kHz ultrasonic vocalization playback at 8 weeks of age (see below).

Surgery, transmitter implantation, heart-rate registration

Radiotelemetric transmitters (HD-S10, Data Sciences International, St. Paul, MN, USA) were implanted into the abdominal aorta of rats in specified groups as previously described (Olszyński et al., 2020, Olszyński et al., 2021). An illustrative image with the surgery details can be found elsewhere (Figure 5 in Pestana-Oliveira et al., 2020; please note, tissue glue was used instead of cellulose patches and silk sutures). The signal was collected by receivers (RSC-1, Data Sciences International, St. Paul, MN, USA) as previously described (Olszyński et al., 2020, Olszyński et al., 2021, Olszyński et al., 2022). Readings were processed using Dataquest ART (version 4.36, Data Sciences International) for trace fear conditioning (Tab. 1/Exp. 1) and Ponemah (version 6.32, Data Sciences International) software for other experiments (Tab. 1/Exp. 2-3 and playback experiments).

Fear conditioning

All conditioning procedures were conducted in a chamber (VFC-008-LP, Med Associates, Fairfax, VT, USA) located in an outer cubicle (MED-VFC2-USB-R, Med Associates) equipped with an ultrasound CM16/CMPA condenser microphone (Avisoft Bioacoustics, Berlin, Germany). Ultrasonic vocalizations were recorded via Avisoft USGH Recorder (Avisoft Bioacoustics), and rat behavior was recorded via NIR monochrome camera (VID-CAM-MONO-6, Med Associates). All procedures were described in detail before (Olszyński et al., 2021, Olszyński et al., 2022).

Added in the revised version: Trace fear conditioning

Earlier version: Trace fear conditioning (Tab. 1/Exp. 1/#1-4, n = 34 rats) was performed similarly to some previous reports (e.g., Jahołkowski et al., 2009). Rats were individually placed in the fear conditioning apparatus in one of two different contexts: A (safe) or B (unsafe). Context A was in an illuminated room with the cage interior with white light, the cage floor was made of solid plastic, and the cage was scented with lemon odor, cleaned with a 10% ethanol solution; the experimenter was male wearing white gloves. Context B was a different, dark room, with the cage interior with green light, the floor was made of metal bars, and the cage was scented with mint odor, cleaned with 1% acetic acid; the experimenter was female with violet gloves. The procedure: on day -2, each rat was habituated to context A for 20 min; on day -1, habituated to context B for 20 min; on day 0, each rat was placed for 52 min in context A; on day 1, after 10 min in context B, the rat received 10 conditioning stimuli (15-s-long sine wave tone, 5 kHz, 85 dB) followed by a 30 s trace period and a foot-shock (1 s, 1 mA) and 210 s inter-trial interval, i.e., ITI; total session duration: 52 min. Control rats were subjected to the same procedures but did not receive the electric shock at the end of trace periods. The animals were tested with the same protocol without shocks in context A (day 2) and context B (day 3). During the test session, control animals showed a lower level of freezing than conditioned animals (1.3 ± 0.8% vs. 19.7 ± 4.3% during the first 5 min in unsafe context B and 0.4 ± 0.3% vs. 9.9 ± 1.9% during 10 s following the time of expected shock in context B, results averaged from the first 3 out of 10 trials; p = 0.0003 and p = 0.0001, respectively, Mann-Whitney); none of the control animals emitted 44-kHz calls, neither the fear conditioning day nor the test days.

Revised version: (Tab. 1/Exp. 1/#1-4, n = 34 rats) was performed similarly to some previous reports (e.g., Jahołkowski et al., 2009). Rats were individually placed in the fear conditioning apparatus in one of two different contexts: A (safe) or B (unsafe). Context A was in an illuminated room with the cage interior with white light, the cage floor was made of solid plastic, and the cage was scented with lemon odor, cleaned with a 10% ethanol solution; the experimenter was male wearing white gloves. Context B was a different, dark room, with the cage interior with green light, the floor was made of metal bars, and the cage was scented with mint odor, cleaned with 1% acetic acid; the experimenter was female with violet gloves. The procedure: on day -2, each rat was habituated to context A for 20 min; on day -1, habituated to context B for 20 min; on day 0, each rat was placed for 52 min in context A; on day 1, after 10 min in context B, the rat received 10 conditioning stimuli (15-s-long sine wave tone, 5 kHz, 85 dB) followed by a 30 s trace period and a foot-shock (1 s, 1 mA) and 210 s inter-trial interval, i.e., ITI; total session duration: 52 min. Control rats were subjected to the same procedures but did not receive the electric shock at the end of trace periods. The animals were tested with the same protocol without shocks in context A (day 2) and context B (day 3). During the test session, control animals showed a lower level of freezing than conditioned animals (1.3 ± 0.8% vs. 19.7 ± 4.3% during the first 5 min in unsafe context B and 0.4 ± 0.3% vs. 9.9 ± 1.9% during 10 s following the time of expected shock in context B, results averaged from the first 3 out of 10 trials; p = 0.0003 and p = 0.0001, respectively, Mann-Whitney); none of the control animals emitted 44-kHz calls, neither the fear conditioning day nor the test days.

Added in the revised version: Delay fear conditioning

Earlier version: Delay fear conditioning. (Tab. 1/Exp. 2-3/#5-13, n = 184 rats) The procedure and its results were described before (Olszyński et al., 2021, Olszyński et al., 2022); rats received 1, 6 or 10 conditioning stimuli (20-s-long white light co-terminating with an electric foot-shock, 1 s, 1 mA). For control rats, an equal time-length procedure was done for each conditioning protocol, i.e., the same parameters as in 1, 6 or 10 stimuli groups, with no shock. Control animals showed a lower level of freezing than conditioned animals. There were only 4 ultrasonic calls we classified as 44-kHz vocalizations among 4,126 vocalizations emitted by the control rats during training and testing. We did not observe any difference in the number of 44-kHz vocalizations between Wistar rats with transmitters vs. without transmitters during delay conditioning training (p = 0.8642, Mann-Whitney). These two groups were therefore reported together.

Revised version: (Tab. 1/Exp. 2-3/#5-13, n = 184 rats) The procedure and its results were described before (Olszyński et al., 2021, Olszyński et al., 2022); rats received 1, 6 or 10 conditioning stimuli (20-s-long white light co-terminating with an electric foot-shock, 1 s, 1 mA). For control rats, an equal time-length procedure was done for each conditioning protocol, i.e., the same parameters as in 1, 6 or 10 stimuli groups, with no shock. Control animals showed a lower level of freezing than conditioned animals. There were only 4 ultrasonic calls we classified as 44-kHz vocalizations among 4,126 vocalizations emitted by the control rats during training and testing. We did not observe any difference in the number of 44-kHz vocalizations between Wistar rats with transmitters vs. without transmitters during delay conditioning training (p = 0.8642, Mann-Whitney). These two groups were therefore reported together.

Measuring freezing

Freezing behavior was scored automatically using Video Freeze software (Med Associates) with a default motion index threshold of 18. To avoid including brief moments of the animal’s stillness, freezing was measured only if the animal did not move for at least 1 s, i.e., 30 video frames, with some exceptions, see next.

Vocalization-nested freezing behavior

Freezing at the exact times of ultrasonic calling was measured in rats that had undergone 10 trials of fear conditioning which produced 44-kHz calls (n = 32, selected from Tab. 1/Exp. 1-3/#2,4,8,13). From each rat, one 44-kHz call was randomly selected along with the long 22-kHz call closest to it. Such pairs of vocalizations were selected with either ≥150 ms duration (n = 32) or ≥500 ms duration (n = 28). For each pair of vocalizations, the freezing behavior was calculated from the entire duration of the shorter call and for the equal-time-length period in the middle of the longer vocalization. Due to the shortened time-scale, the minimal freezing detection window was reduced to 3 frames for ≥150-ms-long calls as well as 5, 10, and 15 frames – for ≥500-ms calls.

Ultrasonic playback

It was performed as described previously (Olszyński et al., 2020, Olszyński et al., 2021, Olszyński et al., 2022) in individual experimental cages with acoustic stimuli presented through a Vifa ultrasonic speaker (Avisoft Bioacoustics, Berlin, Germany) connected to an UltraSoundGate Player 116 (Avisoft Bioacoustics). Ultrasonic vocalizations emitted by the rat were recorded by a CM16/CMPA condenser microphone (Avisoft Bioacoustics). Both playback and recording of calls were performed using Avisoft Recorder USGH software (version 4.2.28, Avisoft Bioacoustics). The locomotor activity was recorded with an acA1300-60gc camera (Basler AG, Ahrensburg, Germany). There were 8 sets of ultrasonic vocalizations presented:

44-kHz long calls, 8 calls in 1 repeat, constant frequency (2.7 ± 0.1 kHz max-min frequency difference), 42.1 ± 0.2 kHz peak frequency, 1064.3 ± 89.6 ms duration with 199.0 ± 14.7 ms sound intervals;

22-kHz long calls, 8 calls in 1 repeat, typical long 22-kHz vocalizations, constant frequency (1.9 ± 0.9 kHz max-min frequency difference), 24.5 ± 0.2 kHz peak frequency, 1066.4 ± 90.2 ms duration with 195.6 ± 15.5 ms sound intervals;

22-kHz short modulated calls, 26 calls in 2 repeats, short (<300 ms), not resembling typical 22-kHz long calls (5.3 ± 0.4 kHz max-min frequency difference), 22.7 ± 0.6 kHz peak frequency, 24.7 ± 1.6 ms duration with 172.8 ± 5.6 ms sound intervals;

22-kHz short flat calls, 43 calls in 1 repeat, short (<300 ms), resembling typical 22-kHz long calls, constant frequency (2.3 ± 0.1 kHz max-min frequency difference), 25.1 ± 0.3 kHz peak frequency, 102.4 ± 10.9 ms duration with 132.1 ± 6.2 ms sound intervals;

50-kHz modulated calls, 23 calls in 2 repeats, moderately modulated (8.6 ± 0.3 kHz max-min frequency difference), 61.0 ± 0.8 kHz peak frequency, 37.6 ± 1.5 ms duration with 183.7 ± 4.5 ms sound intervals;

50-kHz flat calls, 29 calls in 2 repeats, constant frequency (4.2 ± 0.2 kHz max-min frequency difference), 53.5 ± 0.5 kHz peak frequency, 66.2 ± 3.8 ms duration with 144.1 ± 4.4 ms sound intervals;

50-kHz trill calls, 29 calls in 2 repeats, highly modulated (37.4 ± 1.7 kHz max-min frequency difference), 68.0 ± 0.9 kHz peak frequency, 53.7 ± 1.4 ms duration with 158.5 ± 4.9 ms sound intervals;

50-kHz kHz mixed calls, used previously in Olszyński et al. (2020), Olszyński et al. (2021), and Olszyński et al. (2022), 28 calls, in 3 repeats, frequency modulated and trill subtypes, 9.8 ± 1.9 kHz max-min frequency difference, 58.6 ± 0.7 kHz peak frequency, 28.4 ± 1.6 ms duration with 91.4 ± 1.4 ms sound intervals.

Calls were presented with a sampling rate of 250 kHz in 16-bit format. All calls except for 50-kHz mixed calls were collected in our laboratory from fear conditioning or playback experiments. Calls in the same set were taken from one animal wherever possible. The sound interval was adjusted if it was peculiarly long or the sequence was interrupted by other types of calls in the original recordings.

Added in the revised version: Playback procedure, rats with transmitters

Earlier version: Playback procedure, rats with transmitters; as previously described (Olszyński et al., 2020, Olszyński et al., 2021, Olszyński et al., 2022). Before playback presentation, animals were habituated for 3 min to the experimental conditions, i.e., recording cage, presence of the speaker and microphone, over 4 days. Habituated rats then underwent a playback procedure, in short, after 10 min of silence, the rats were exposed to four 10-s-long call sets (either aversive or appetitive) with 5-min-long ITI in-between; a rat that received appetitive playback was followed by a rat receiving aversive playbacks etc. Also, the order of the presented sets was randomized between animals. The aversive-calls playback contained sets nos. 1-4. The appetitive-calls playback contained sets nos. 5-8. Since initial analysis showed no differences within responses to 22-kHz aversive sets and within responses to 50-kHz appetitive sets, we decided to show the results following playback of 44-kHz long calls (set no. 1), 22-kHz long calls (set no. 2), and 50-kHz modulated calls (set no. 5) only.

Revised version: As previously described (Olszyński et al., 2020, Olszyński et al., 2021, Olszyński et al., 2022). Before playback presentation, animals were habituated for 3 min to the experimental conditions, i.e., recording cage, presence of the speaker and microphone, over 4 days. Habituated rats then underwent a playback procedure, in short, after 10 min of silence, the rats were exposed to four 10-s-long call sets (either aversive or appetitive) with 5-min-long ITI in-between; a rat that received appetitive playback was followed by a rat receiving aversive playbacks etc. Also, the order of the presented sets was randomized between animals. The aversive-calls playback contained sets nos. 1-4. The appetitive-calls playback contained sets nos. 5-8. Since initial analysis showed no differences within responses to 22-kHz aversive sets and within responses to 50-kHz appetitive sets, we decided to show the results following playback of 44-kHz long calls (set no. 1), 22-kHz long calls (set no. 2), and 50-kHz modulated calls (set no. 5) only.

Added in the revised version: Playback procedure, rats without transmitters

Earlier version: Playback procedure, rats without transmitters. Before playback presentation, animals were habituated for 3 min to the experimental conditions, i.e., recording cage, presence of the speaker and microphone, over 4 days. After 5 min of initial silence, the rats were presented with two 10-s-long playback sets of either 22-kHz (set no. 2; n = 8) or 44-kHz calls (set no. 1; n = 8), followed by one 50-kHz modulated call 10-s set (no. 5) and another two playback sets of either 44-kHz or 22-kHz calls not previously heard. The playback presentations were separated by 3 min ITI. Responses to the pairs of playback sets were averaged.

Revised version: Before playback presentation, animals were habituated for 3 min to the experimental conditions, i.e., recording cage, presence of the speaker and microphone, over 4 days. After 5 min of initial silence, the rats were presented with two 10-s-long playback sets of either 22-kHz (set no. 2; n = 8) or 44-kHz calls (set no. 1; n = 8), followed by one 50-kHz modulated call 10-s set (no. 5) and another two playback sets of either 44-kHz or 22-kHz calls not previously heard. The playback presentations were separated by 3 min ITI. Responses to the pairs of playback sets were averaged.

Added in the revised version: Locomotor activity in playback

Earlier version: Locomotor activity in playback. An automated video tracking system (Ethovision XT 10, Noldus, Wageningen, The Netherlands) was used to measure the total distance travelled (cm). Proximity to the speaker was expressed as the percentage of time spent in the half of the cage closer to the ultrasonic speaker. Center-point of each animal’s shape was used as a reference point for measurements of locomotor activity thus registering only full-body movements.

Revised version: An automated video tracking system (Ethovision XT 10, Noldus, Wageningen, The Netherlands) was used to measure the total distance travelled (cm). Proximity to the speaker was expressed as the percentage of time spent in the half of the cage closer to the ultrasonic speaker. Center-point of each animal’s shape was used as a reference point for measurements of locomotor activity thus registering only full-body movements.

Analysis of ultrasonic vocalizations

Audio recordings were analyzed manually using SASLab Pro (version 5.2.xx, Avisoft Bioacoustics) as described (Olszyński et al., 2020, Olszyński et al., 2021, Olszyński et al., 2022) to measure key features of calls and categorize them into subtypes.

Sound mean power was measured as the average spectra power density of the vocalization contour using DeepSqueak software. Initially, calls were detected using the default rat long-vocalization neural network (Long Rat Detector YOLO R1) and subsequently manually reviewed and corrected where necessary. We analyzed a subset of Wistar rats subjected to 10-trial fear conditioning that emitted more than 20 instances of 44-kHz calls during the fear conditioning session (n = 17, selected from Tab. 1/Exp. 1-3/#2,4,8,13). It is important to note that due to the directional characteristics of the microphones used, angular attenuation occurred during audio recording. This phenomenon results in a selective reduction in the intensity of higher frequency sounds, dependent on the angle between the sound emitter and the microphone (as specified in the CM16/CMPA microphone hardware specification page, Avisoft Bioacoustics website). In our experimental setup, we approximated a 45° angle between the plane of the rat’s head and the plane of the microphone’s membrane. This angle corresponds to an estimated 10 dB attenuation (adopting a conservative estimate) of 40-kHz frequencies compared to 20-kHz frequencies for which there is even a small dB gain due to these hardware properties, 44-kHz calls are predicted to be approximately at least 10 dB louder in reality than what was recorded.

Added in the revised version: 22-kHz vs. 44-kHz frequency ratio

Earlier version: 22-kHz vs. 44-kHz frequency ratio. A clear transition point between 22-kHz and 44-kHz long calls was observed in n = 13 Wistar rats and n = 1 SHR. In each case, ten 22-kHz calls followed by ten 44-kHz calls were analyzed (n = 14, selected from Tab. 1/Exp. 1-3/#2,4,6-8,10-13).

Revised version: A clear transition point between 22-kHz and 44-kHz long calls was observed in n = 13 Wistar rats and n = 1 SHR. In each case, ten 22-kHz calls followed by ten 44-kHz calls were analyzed (n = 14, selected from Tab. 1/Exp. 1-3/#2,4,6-8,10-13).

Added in the revised version: Step up and step down frequency ratio

Earlier version: Step up and step down frequency ratio. Rats which emitted at least five vocalizations of the specific subtype were analyzed (step up, n = 14; step down, n = 13; selected from Tab. 1/Exp. 1-3/#2,4,7,8,13; 5 calls of the two subtypes from each rat were chosen randomly and the frequencies of their elements were measured.

Revised version: Rats which emitted at least five vocalizations of the specific subtype were analyzed (step up, n = 14; step down, n = 13; selected from Tab. 1/Exp. 1-3/#2,4,7,8,13; 5 calls of the two subtypes from each rat were chosen randomly and the frequencies of their elements were measured.

Ultrasonic vocalizations clustering (two independent methods)

Calls of conditioned and control animals were taken from all fear conditioning training sessions (Tab. 1/Exp. 1-3, n = 218). We used DBSCAN algorithm (Ester et al., 1996); a density based method, from the scikit-learn (sklearn) Python package, because of its ability to detect a desired number of clusters of arbitrary shape; with two main input parameters: MinPts (minimal number of points forming the core of the cluster) and ε (the maximum distance two points can be from one another while still belonging to the same cluster). To avoid detecting small clusters, we limited MinPts to 150 samples. The heuristic method described by Ester et al. (Ester et al., 1996) was implemented to find the initial range of ε. All the input data were standardized. The silhouette coefficient (Rousseeuw, 1987) was used to control the quality of the clustering. Maximizing ε among different ranges helped to select the most relevant number of identified clusters. Clustering with ε in the range of 0.14–0.2 resulted in a silhouette coefficient around 0.2–0.5.

Added in the revised version: K-means algorithm

Earlier version: K-means algorithm. Vocalizations of selected fear-conditioned rats with 6-10 shocks and >30 of 44-kHz calls (n = 26, selected from Tab. 1/Exp. 1-3/#2,4,7,8,11-13) were detected using a built-in neural network for long rat calls (Long Rat Detector YOLO R1) on DeepSqueak (Coffey et al., 2019) software (version 3.0.4) running under MATLAB (version 2021b, MathWorks, Natick, MA, USA) and manually revised for missed and mismatched calls. Unsupervised k-means clustering was based on call contour, frequency and duration variables, with equal weights assigned, and several descending elbow optimization parameters were used to obtain different maximum numbers of clusters together with Uniform Manifold Approximation and Projection for Dimension Reduction (UMAP) (McInnes et al., 2018) for superimposing and visualization of clusters.

Revised version: Vocalizations of selected fear-conditioned rats with 6-10 shocks and >30 of 44-kHz calls (n = 26, selected from Tab. 1/Exp. 1-3/#2,4,7,8,11-13) were detected using a built-in neural network for long rat calls (Long Rat Detector YOLO R1) on DeepSqueak (Coffey et al., 2019) software (version 3.0.4) running under MATLAB (version 2021b, MathWorks, Natick, MA, USA) and manually revised for missed and mismatched calls. Unsupervised k-means clustering was based on call contour, frequency and duration variables, with equal weights assigned, and several descending elbow optimization parameters were used to obtain different maximum numbers of clusters together with Uniform Manifold Approximation and Projection for Dimension Reduction (UMAP) (McInnes et al., 2018) for superimposing and visualization of clusters.

Quantification and statistical analysis

Earlier version: Data were analyzed using non-parametric Friedman, Wilcoxon, Mann-Whitney tests with GraphPad Prism 8.4.3 (GraphPad Software, San Diego, CA, USA); the p values are given, p < 0.05 as the minimal level of significance. Figures were prepared using the same software and depict average values with a standard error of the mean (SEM).

Revised version: Data were analyzed using non-parametric Friedman, Wilcoxon, Mann-Whitney tests with GraphPad Prism 8.4.3 (GraphPad Software, San Diego, CA, USA); the p values are given, p < 0.05 as the minimal level of significance. In particular, the Friedman test was used to assess the presence of change within the sequence of several ITI, while the Wilcoxon test was used for the difference between the first and the last ITI analyzed. Figures were prepared using the same software and depict average values with a standard error of the mean (SEM).

Added in the revised version: Data availability

Added in the revised version: Raw data (calls’ peak frequency and duration) analyzed, ultrasonic playback files used (.wav), data supporting clustering files for DBSCAN (.csv), and extracted call contours for k-means (.mat) have been deposited to Mendeley Data at http:/to be provided/. The other data in this study are available from the corresponding author upon request.

Added in the revised version: Supporting information